<problem display_name="Primes in Names" showanswer="always" weight="10" attempts="">

  <p>Variable names can now include primes at the end! Here, the answer is "x'+v*t" (a Galilean change of coordinate).</p>

<script type="text/python" system_path="python_lib">
from mitxgraders import *
grader = FormulaGrader(
    answers="x'+v*t",
    variables=["x'", 'v', 't']
)
</script>

<p style="display:inline">[mathjaxinline]x =[/mathjaxinline] </p>
<customresponse cfn="grader" inline="1">
  <textline math="1" inline="1" correct_answer="x'+v*t"/>
</customresponse>

<a href="https://github.com/mitodl/mitx-grading-library/tree/master/course/problem/formula4.xml" target="_blank">View source</a>

</problem>
