<problem display_name="User Functions and Constants" showanswer="always" weight="10" attempts="">

<p>It is possible to specify user-defined constants and functions. Here, we add a constant tau (= 2*pi), and also define the Heaviside step function H(x). You can also specify an unknown function, for which we sample a well-behaved random function. In the following example, the answer is <code>H(x)*f''(tau)</code>, where f'' is a random function.</p>

<script type="text/python" system_path="python_lib">
from mitxgraders import *

def heaviside(x):
    if x > 0:
        return 1
    return 0

grader = FormulaGrader(
    user_constants={'tau': 6.2831853072},
    user_functions={
        "H": heaviside,
        "f''": RandomFunction()
    },
    variables=['x'],
    sample_from={
        'x': RealInterval([-5,5])
    }
)
</script>

<customresponse cfn="grader" inline="1" expect="H(x)*f''(tau)">
  <textline math="1" inline="1"/>
</customresponse>

<p>Try typing <code>1/H(x)</code> in the box above. Do you see a problem? Try the box below - it's exactly the same, except that we've loaded our AsciiMath preprocessor to ensure that the rendering is picture perfect. We strongly recommend using our AsciiMath renderer to ensure that students see a math preview that matches with what they should be writing on paper.</p>

<customresponse cfn="grader" inline="1" expect="H(x)*f''(tau)">
  <textline math="1" inline="1" preprocessorClassName="MJxPrep" preprocessorSrc="/static/MJxPrep.js"/>
</customresponse>

<p><a href="https://github.com/mitodl/mitx-grading-library/tree/master/course/problem/formula5.xml" target="_blank">View source</a></p>

</problem>
