<problem display_name="User Functions and Constants" showanswer="always" weight="10" attempts="">

  <p>It is possible to specify user-defined constants and functions. Here, we add a constant tau (= 2*pi), and also define the Heaviside step function H(x). You can also specify an unknown function, for which we sample a well-behaved random function. In the following example, the answer is "H(x)*f''(tau)", where f'' is a random function.</p>

<script type="text/python" system_path="python_lib">
from mitxgraders import *

def heaviside(x):
    if x > 0:
        return 1
    return 0

grader = FormulaGrader(
    answers="H(x)*f''(tau)",
    user_constants={'tau': 6.2831853072},
    user_functions={
        "H": heaviside,
        "f''": RandomFunction()
    },
    variables=['x'],
    sample_from={
        'x': RealInterval([-5,5])
    }
)
</script>

<customresponse cfn="grader" inline="1">
  <textline math="1" inline="1" correct_answer="H(x)*f''(tau)"/>
</customresponse>

<p>User-defined functions often don't render quite right when using textline inputs. For example, try typing in "1/H(x)" in the box above. If you make extensive use of user-defined functions, we suggest incorporating those functions into the AsciiMath renderer definitions file we provide. See the documentation for details.</p>

<a href="https://github.com/mitodl/mitx-grading-library/tree/master/course/problem/formula5.xml" target="_blank">View source</a>

</problem>
