<problem display_name="Improvements over edX" showanswer="always" weight="10" attempts="">

  <p>Fixes to edX (unmatched parentheses, case on functions, variable as a function, 2m, imaginary values)</p>

<script type="text/python" system_path="python_lib">
from mitxgraders import *
grader = StringGrader(
    answers='cat'
)
</script>

<customresponse cfn="grader">
  <textline correct_answer="cat"/>
</customresponse>

<a href="https://github.com/mitodl/mitx-grading-library/tree/master/course/problem/formula1.xml" target="_blank">View source</a>

</problem>
