<problem display_name="Improvements over edX" showanswer="always" weight="10" attempts="">

  <p>The following is a simple formularesponse problem, except that we're using our library instead of edX. There is a variable "m", and the answer is "m*(1-sin(2*m))".</p>

  <p>The interesting part is what happens when you have a mistake in your input. Try the following:</p>

  <ul>
    <li>m*(1-sin(2*m) (missing closing parentheses)</li>
    <li>m*(1-sin(2*m))) (extra closing parentheses)</li>
    <li>m*(1-Sin(2*m)) (wrong capitalization for function name)</li>
    <li>m(1-sin(2*m)) (forgot multiplication before brackets)</li>
    <li>m*(1-sin(2m)) (forgot multiplication in 2*m)</li>
    <li>x*(1-sin(2*m)) (unknown variable x)</li>
    <li>1/0 (division by zero)</li>
    <li>10^400 (overflow)</li>
  </ul>

<script type="text/python" system_path="python_lib">
from mitxgraders import *
grader = FormulaGrader(
    answers='m*(1-sin(2*m))',
    variables=['m']
)
</script>

<customresponse cfn="grader">
  <textline math="1" correct_answer="m*(1-sin(2*m))"/>
</customresponse>

<p>For comparison, the following box is an edX formularesponse grader with exactly the same configuration.</p>

<formularesponse type="cs" samples="m@1:3#5" answer="m*(1-sin(2*m))">
  <textline math="1">
    <responseparam type="tolerance" default="0.1%"/>
  </textline>
</formularesponse>


<a href="https://github.com/mitodl/mitx-grading-library/tree/master/course/problem/formula1.xml" target="_blank">View source</a>

</problem>
