<problem display_name="Partially-Ordered Input" showanswer="always" weight="10" attempts="">

<p>Sometimes you will want partially-ordered input. In the following example, we want "Halloween" in the first box, and "ghost", "bat" and "pumpkin" in any order in the other boxes.</p>

<script type="text/python">
from mitxgraders import *
grader = ListGrader(
    answers=[
        'Halloween',
        ['bat', 'ghost', 'pumpkin']
    ],
    subgraders=[
        StringGrader(),
        ListGrader(
            subgraders=StringGrader()
        )
    ],
    ordered=True,
    grouping=[1, 2, 2, 2]
)
</script>

<customresponse cfn="grader" inline="1">
  <p style="display:inline">Event: </p><textline inline="1" correct_answer="Halloween"/>
  <br/>
  <p style="display:inline">Object: </p><textline inline="1" correct_answer="bat"/>
  <br/>
  <p style="display:inline">Object: </p><textline inline="1" correct_answer="ghost"/>
  <br/>
  <p style="display:inline">Object: </p><textline inline="1" correct_answer="pumpkin"/>
</customresponse>

<p>Note that repeating an object will be graded as incorrect.</p>

<p><a href="https://github.com/mitodl/mitx-grading-library/tree/master/course/problem/list4.xml" target="_blank">View source</a></p>

</problem>
