<problem display_name="Grouped Input" showanswer="always" weight="10" attempts="">

  <p>Often you will want to group input. There are three lions, two dogs and one bird. In the following boxes, write the singular of each animal, and the corresponding number of each animal. Order is unimportant, but the pairing must be right.</p>

<script type="text/python" system_path="python_lib">
from mitxgraders import *
grader = ListGrader(
    answers=[
        ['bird', '1'],
        ['dog', '2'],
        ['lion', '3']
    ],
    subgraders=ListGrader(
        subgraders=[StringGrader(), NumericalGrader()],
        ordered=True
    ),
    grouping=[1, 1, 2, 2, 3, 3]
)
</script>

<customresponse cfn="grader" inline="1">
  <table>
    <tr>
      <th>Animal</th>
      <th>Number</th>
    </tr>
    <tr>
      <td><textline correct_answer="dog"/></td>
      <td><textline correct_answer="2"/></td>
    </tr>
    <tr>
      <td><textline correct_answer="bird"/></td>
      <td><textline inline="1" correct_answer="1"/></td>
    </tr>
    <tr>
      <td><textline inline="1" correct_answer="lion"/></td>
      <td><textline inline="1" correct_answer="3"/></td>
    </tr>
  </table>
</customresponse>

<p>Note that such questions are often most cleanly presented using an HTML table, as shown here.</p>

<a href="https://github.com/mitodl/mitx-grading-library/tree/master/course/problem/list3.xml" target="_blank">View source</a>

</problem>
