<problem display_name="Grouped Input" showanswer="always" weight="10" attempts="">

  <p>Often you will want to group input. There are three lions, two dogs and one bird. In the following boxes, write the singular of each animal, and the corresponding number of each animal. Order is unimportant, but the pairing must be right.</p>

<script type="text/python" system_path="python_lib">
from mitxgraders import *
grader = ListGrader(
    answers=[
        ['bird', '1'],
        ['dog', '2'],
        ['lion', '3']
    ],
    subgraders=ListGrader(
        subgraders=[StringGrader(), NumericalGrader()],
        ordered=True
    ),
    grouping=[1, 1, 2, 2, 3, 3]
)
</script>

<customresponse cfn="grader" inline="1">
  <p style="display:inline">First animal: </p><textline inline="1" correct_answer="dog"/>
  <p style="display:inline">Number: </p><textline inline="1" correct_answer="2"/>
  <br/>
  <p style="display:inline">Second animal: </p><textline inline="1" correct_answer="bird"/>
  <p style="display:inline">Number: </p><textline inline="1" correct_answer="1"/>
  <br/>
  <p style="display:inline">Third animal: </p><textline inline="1" correct_answer="lion"/>
  <p style="display:inline">Number: </p><textline inline="1" correct_answer="3"/>
</customresponse>


<a href="https://github.com/mitodl/mitx-grading-library/tree/master/course/problem/list3.xml" target="_blank">View source</a>

</problem>
