<problem display_name="Tech Demo: specify_domain" showanswer="always" weight="10" attempts="">

<script type="text/python" system_path="python_lib">
from mitxgraders import MatrixGrader, specify_domain, RealVectors, RealMatrices

# This is a contrived example specify_domain, just to demonstrate its usage
@specify_domain(input_shapes=[1, [2, 3], [3]], display_name='f')
def my_func(a, b, c):
    return a*b*c

# Define a grader using rotate
grader = MatrixGrader(
   answers='f(a, b, c)',
   variables=['a', 'b', 'c'],
   sample_from={
       'b': RealMatrices(shape=[2, 3]),
       'c': RealVectors(shape=3),
   },
   user_functions={
       'f': my_func
   }
)

</script>

<p>If you need to define an author-specified function in a MatrixGrader problem, you can use the <code>specify_domain</code> decorator to provide informative error messages to students when things go wrong. (For example, if a student inputs a scalar to a function whose domain is vectors).</p>

<p>This problem is intended only to show off the syntax (see source code below) and error messages provided by <code>specify_domain</code>.</p>

<p>Available variables and functions:</p>
<ul>
    <li><code>func(arg1, arg2, arg3)</code> expects a scalar <code>arg1</code>, a 3 by 2 matrix <code>arg2</code>, and a 3-component vector <code>arg3</code>.</li>
    <li><code>a</code> is a scalar</li>
    <li><code>b</code> is a 3 by 2 matrix</li>
    <li><code>c</code> is a 3-component vector</li>
</ul>

<p>Suggested inputs:</p>
<ul>
    <li><code>f(a, b, c)</code> is the correct answer</li>
    <li><code>f(a, [1, 2], b)</code>, or anything where arguments have the wrong shape, will produce interesting error messages</li>
</ul>

<customresponse cfn="grader" expect="f(a,b,c)">
  <textline math="true" size="40" />
</customresponse>

<h3>Resources</h3>
<ul>
  <li>
    <a href="https://github.com/mitodl/mitx-grading-library/tree/master/course/problem/matrix6.xml" target="_blank">View source</a>
  </li>
  <li>
    <a href="https://mitodl.github.io/mitx-grading-library/grading_math/functions/#the-specify_domain-decorator" target="_blank">Documentation for <code>specify_domain</code></a>
  </li>
</ul>

</problem>
