<problem display_name="Length Checking" showanswer="always" weight="10" attempts="">

  <p>By default, the grader accepts any number of inputs in the list (try it in some of the other examples!). If you want to raise an error if the number of inputs is wrong, then you can do that too.</p>

  <p>In 3D, what are the components of the \(\hat{y}\) unit vector? (Answer: 0, 1, 0). Try putting in two or four components instead of three.</p>

<script type="text/python" system_path="python_lib">
from mitxgraders import *
mygrader = SingleListGrader(
    answers=['0', '1', '0'],
    subgrader=FormulaGrader(),
    length_error=True
)
</script>

<customresponse cfn="grader">
  <textline correct_answer="0, 1, 0"/>
</customresponse>

<a href="https://github.com/mitodl/mitx-grading-library/tree/master/course/problem/singlelist3.xml" target="_blank">View source</a>

</problem>
