<problem display_name="Multiple Answers" showanswer="always" weight="10" attempts="">

  <p>Here is an example where multiple answers are accepted. You can either enter ("cat" or "feline") and "dog", or "goat" and "vole". The entries should be comma separated, and can appear in any order. Note that mixing one set of answers with the other set will only give you partial credit though.</p>

<script type="text/python" system_path="python_lib">
from mitxgraders import *
grader = SingleListGrader(
    answers=(
        [('cat', 'feline'), 'dog'],
        ['goat', 'vole'],
    ),
    subgrader=StringGrader()
)
</script>

<customresponse cfn="grader">
  <textline correct_answer="cat, dog"/>
</customresponse>

<a href="https://github.com/mitodl/mitx-grading-library/tree/master/course/problem/singlelist2.xml" target="_blank">View source</a>

</problem>
