<problem display_name="Accepting Anything" showanswer="always" weight="10" attempts="">

<p>Sometimes, you may want to accept anything that a student types in (possibly with some caveats, which we'll get to in a bit). To do this, we can use the <code>accept_any</code> option. Here is an example.</p>

<p>Please enter your favorite color in the textbox below.</p>

<script type="text/python">
from mitxgraders import *
grader = StringGrader(
    # Note that the expect entry here is irrelevant
    answers={'expect': '', 'grade_decimal': 1, 'msg': 'Your answer has been recorded.'},
    accept_any=True
)
</script>

<customresponse cfn="grader">
  <textline />
</customresponse>

<p>You can also require the student to enter a minimum number of characters and/or words. The following problem needs at least 10 characters and 3 words. See what happens when you don't meet these requirements too!</p>

<p>What do you think about this problem?</p>

<script type="text/python">
grader2 = StringGrader(
    # Note that the expect entry here is irrelevant
    answers={'expect': '', 'grade_decimal': 1, 'msg': 'Your answer has been recorded.'},
    accept_any=True,
    min_length=10,
    min_words=3,
    explain_minimums='err'
)
</script>

<customresponse cfn="grader2">
  <textline />
</customresponse>

<p><a href="https://github.com/mitodl/mitx-grading-library/tree/master/course/problem/string3.xml" target="_blank">View source</a></p>

</problem>
