<problem display_name="Even and Odd Sums" showanswer="always" weight="10" attempts="">

<p>You can specify that a sum is over even or odd integers only.</p>

<p>Construct a Taylor series for \(\sin(x)\), summing over odd integers only.</p>

<style>
  .xmodule_display.xmodule_ProblemBlock .problem .capa_inputtype.textline input {
    min-width: 0 !important;
  }
  .xmodule_display.xmodule_ProblemBlock div.problem section div span.MathJax {
    display: inline-block !important;
  }
  .xmodule_display.xmodule_ProblemBlock div.problem section div span.MathJax_Preview {
    display: inline-block !important;
  }
</style>

<script type="text/python">
from mitxgraders import *
grader = SumGrader(
    answers={
        'lower':'1',
        'upper':'infty',
        'summand':'i^(n-1)*x^n/fact(n)',
        'summation_variable':'n'
    },
    input_positions = {
        'upper': 1,
        'summand': 2,
        'lower': 3
    },
    variables=['x'],
    sample_from={'x': [0, 0.5]},
    even_odd=1
)
</script>

<span>
<customresponse cfn="grader">
  <table>
    <col style="width:10%"/>
    <col style="width:90%"/>
    <tbody>
      <tr>
        <td colspan="2">
          <textline size="5" correct_answer="infty"/>
        </td>
      </tr>
      <tr>
        <td>
          <p> \( \displaystyle \huge{ \sum }\)</p>
        </td>
        <td>
          <br/>
          <textline inline="1" size="20" correct_answer="i^(n-1)*x^n/fact(n)" math="1" preprocessorClassName="MJxPrep" preprocessorSrc="/static/MJxPrep.js"/>
        </td>
      </tr>
      <tr>
        <td colspan="2">
            <p style="display: inline;">[mathjaxinline]n \text{ odd,} \ n = [/mathjaxinline] </p><textline size="5" correct_answer="1" inline="1"/>
        </td>
      </tr>
    </tbody>
  </table>
</customresponse>
</span>

<p>Here, the correct answer is <code>infty</code>, <code>i^(n-1)*x^n/fact(n)</code> and <code>1</code> (top to bottom), although there are other ways of writing this also.</p>

<p>Unfortunately, students are unable to specify that the sum should be for even/odd integers only; this must be specified by the problem author upon creation.</p>

<p><a href="https://github.com/mitodl/mitx-grading-library/tree/master/course/problem/sum4.xml" target="_blank">View source</a></p>

</problem>
